```python
# Initialize Otter
import otter
grader = otter.Notebook("demo.ipynb")
```

Pregunta 1

<!--
BEGIN QUESTION
name: q1
points: 1
-->

```python
hola_mundo = 'Hola mundo'

```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: Good job!
q1 - 2 message: Good job!
```

```python

```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!
q1 - 1 message: Good job!
q1 - 2 message: Good job!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Output:
```
There was an error generating your LaTeX; showing full error message:
========================================================================================================================
This is BibTeX, Version 0.99d (MiKTeX 22.3)
The top-level auxiliary file: notebook.aux
I found no \citation commands---while reading file notebook.aux
I found no \bibdata command---while reading file notebook.aux
I found no \bibstyle command---while reading file notebook.aux
(There were 3 error messages)
bibtex: major issue: So far, you have not checked for MiKTeX updates.

========================================================================================================================
If the error above is related to xeCJK or fandol in LaTeX and you don't require this functionality, try running again with no_xecjk set to True or the --no-xecjk flag.
C:\Users\joaro\anaconda3\lib\site-packages\otter\check\notebook.py:328: UserWarning: Could not locate a PDF to include
  warnings.warn("
```

Tests for q1 (points: 1)
```
>>> type(hola_mundo) == str
True
```
failure_message: DEBE SER UN STRING
success_message: Good job!
```
>>> 'hola' in hola_mundo.lower()
True
```
failure_message: tienes que decir hola
success_message: Good job!
