```python
# Initialize Otter
import otter
grader = otter.Notebook("demo.ipynb")
```

Pregunta 1


```python
hola_mundo = 'Hola Mundo'
```

```python
grader.check("q_1")
```

Output:
```
q_1 results: All test cases passed!
q_1 - 1 message: Good job!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q_1 (points: None)
```
>>> type(hola_mundo) == str
True
```
success_message: Good job!
